## Problem 9: NYPD Motor Vehicle Collisions

In this notebook you will analyze motor collisions in New York City. This data is taken from [NYC Open Data](https://data.cityofnewyork.us/Public-Safety/NYPD-Motor-Vehicle-Collisions/h9gi-nx95). This is a breakdown of every collision in NYC by location and injury. Each record represents a collision in NYC by city, borough, precinct and cross street.

This notebook has a total of 4 exercises worth a total of 10 points.



### Setup
Run the following code cells to setup your environment

Please ensure that **Python 3.5** is selected to run this notebook

```python
import pandas as pd
import numpy as np
import sqlite3 as db
```

```python
conn = db.connect('../resource/lib/publicdata/mt2nb9/nypd.db')
c = conn.cursor()
c.execute('SELECT type, name, sql FROM sqlite_master')
results = c.fetchall()
for table in results:
    print(table)
```

Output:
```
('table', 'nypd', 'CREATE TABLE "nypd" (\n"Unique_key" TEXT,\n  "DATE" TEXT,\n  "TIME" TEXT,\n  "BOROUGH" TEXT,\n  "Zip_code" TEXT,\n  "LATITUDE" REAL,\n  "LONGITUDE" REAL,\n  "CONTRIBUTING FACTOR VEHICLE 1" TEXT,\n  "Num_of_persons_injured" INTEGER,\n  "VEHICLE TYPE CODE 1" TEXT\n)')
('table', 'Detail_injured', 'CREATE TABLE "Detail_injured" (\n"Unique_key" TEXT,\n  "Num_of_persons_injured" INTEGER,\n  "Num_of_persons_killed" INTEGER,\n  "Num_of_pedestrians_injured" INTEGER,\n  "Num_of_pedestrians_killed" INTEGER\n)')
```

### Table Details
Please note that there are two tables in the database

<b>1. Table: nypd</b> 

   <b>Columns:</b> Unique_key, DATE, TIME, BOROUGH, Zip_code, LATITUDE, LONGITUDE, CONTRIBUTING FACTOR VEHICLE 1, Num_of_Persons_injured, VEHICLE TYPE CODE 1
   
   
<b>2. Table: Detail_injured </b>

   <b>Columns:</b> Unique_key, Num_of_persons_injured, Num_of_persons_killed, Num_of_pedestrans_injured, Num_of_pedestrians_killed
   
   Table **`Detail_injured`** is only used in Exercise-3.

### Exercises

**Exercise 0** (2 points) Using the table **`nypd`**, here are some tasks to complete in this exercise.
 
1. Using the column, `DATE`, which is in the form mm/dd/yyyy, add a new column to the dataset called `Month`. 
2. Then return the total number of PERSONS INJURED (total_num_of_injured) in each month, by the descending order of Num_of_injured.
 
 Save your result in a table named **`new_dataset`**.
 
 The format of your table should be:


-----| month | total_num_of_injured |
-----|-------|----------------------|
 0   |   10  |  11234               | 



```python
query = '''
    SELECT DATE, Num_of_persons_injured 
     from  nypd
    
'''
df=pd.read_sql_query(query, conn)
df['DATE'] = pd.to_datetime(df['DATE'])
df['month']=df['DATE'].dt.month
df2=df.groupby(['month']).sum()
df2.rename(columns={'Num_of_persons_injured':'total_num_of_injured'},inplace=True)
df2.reset_index(level=0, inplace=True)
df2['month']=df2['month'].map("{:02}".format)
new_dataset=df2
new_dataset.sort_values(by='total_num_of_injured',ascending=False, inplace=True)

print(new_dataset.head(10))
```

Output:
```
month  total_num_of_injured
6     07                 24136
9     10                 23181
7     08                 23132
8     09                 22977
11    12                 22032
10    11                 21679
5     06                 21055
4     05                 20408
0     01                 19096
1     02                 16521
```

```python
new_dataset.sort_values(by='total_num_of_injured',ascending=False, inplace=True)
```

<!-- grade_id: exercise0, points: 2, grade: true, solution: false -->
```python
## Test Cell: Check Month & Total number of injured ##
assert (new_dataset[new_dataset['month']=='06']['total_num_of_injured'] == 21055.0).all()
assert (new_dataset['total_num_of_injured'].sum() == 245979.0)

print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Exercise 1 (2 point):** Using the table **`nypd`**, here are some tasks to complete in this exercise.
There are many contributing factors for the collisions. please list top 10 contributing factors (unspecified is not a factor). Save your result in a table named **`top10_factors`**, by the descending order of cnt.

The format of your table should be:


-----| CONTRIBUTING FACTOR VEHICLE 1   |  cnt  |
-----|---------------------------------|-------|
 0   |             Lost Consciousness  |  11234| 



```python
query = '''
    SELECT * 
    from  nypd 
   
'''
df3=pd.read_sql_query(query, conn)
df3
```

Output:
```
Unique_key        DATE   TIME        BOROUGH Zip_code   LATITUDE  \
0         4090610  03/02/2019   0:00          BRONX    10466  40.893390   
1         4090611  03/02/2019   0:00          BRONX    10469  40.870290   
2         4090424  03/02/2019   0:00       BROOKLYN    11206  40.705185   
3         4089381  03/02/2019   0:00       BROOKLYN    11228  40.622364   
4         4089728  03/02/2019   0:00         QUEENS    11365  40.734680   
5         4089549  03/02/2019   0:00         QUEENS    11372  40.750010   
6         4089855  03/02/2019   0:00         QUEENS    11385  40.709380   
7         4090165  03/02/2019   0:00       BROOKLYN    11236  40.651520   
8         4090229  03/02/2019   0:01       BROOKLYN    11203  40.653683   
9         4089405  03/02/2019   0:01      MANHATTAN    10033  40.848392   
10        4090720  03/02/2019   0:05       BROOKLYN    11208  40.679110   
11        4090556  03/02/2019   0:15         QUEENS    11415  40.713100   
12        4089734  03/02/2019   
```

```python
query = '''
    SELECT * 
    from  nypd 
   
'''
df3=pd.read_sql_query(query, conn)
cols=['CONTRIBUTING FACTOR VEHICLE 1', 'Num_of_persons_injured']
df4=df3[cols]
df5=df4.groupby(['CONTRIBUTING FACTOR VEHICLE 1']).count()
df5.rename(columns={'Num_of_persons_injured':'cnt'},inplace=True)
df5.sort_values(by='cnt',ascending=False, inplace=True)
df5.reset_index(level=0, inplace=True)
top10_factors=df5.iloc[1:11]
```

```python
###
### YOUR CODE HERE
###
cols=['CONTRIBUTING FACTOR VEHICLE 1', 'Num_of_persons_injured']
df4=df3[cols]
df5=df4.groupby(['CONTRIBUTING FACTOR VEHICLE 1']).count()
# df5
df5.rename(columns={'Num_of_persons_injured':'cnt'},inplace=True)
df5.sort_values(by='cnt',ascending=False, inplace=True)
df5.reset_index(level=0, inplace=True)
top10_factors=df5.iloc[1:11]
# Show your solution:
print(top10_factors)


```

Output:
```
CONTRIBUTING FACTOR VEHICLE 1     cnt
1   Driver Inattention/Distraction  166711
2    Failure to Yield Right-of-Way   58626
3                 Backing Unsafely   41497
4                  Other Vehicular   33228
5            Following Too Closely   26373
6                  Fatigued/Drowsy   25211
7               Turning Improperly   23777
8              Passing Too Closely   19249
9   Passing or Lane Usage Improper   18120
10              Lost Consciousness   15152
```

<!-- grade_id: exercise1, points: 2, grade: true, solution: false -->
```python
## Test Cell: Check the contributing factors ##
assert top10_factors['cnt'].sum() == 427944
assert top10_factors['CONTRIBUTING FACTOR VEHICLE 1'].iloc[0] == 'Driver Inattention/Distraction'
assert top10_factors['CONTRIBUTING FACTOR VEHICLE 1'].iloc[9] == 'Lost Consciousness'
assert (top10_factors['CONTRIBUTING FACTOR VEHICLE 1']=='Fatigued/Drowsy').any()

print("\n(Passed!)")
```

Output:
```
(Passed!)
```

**Exercise 2 (3 points):** Using the table **`nypd`**, here are some tasks to complete in this exercise.

There are five boroughs in NYC, please list ONE location in QUEENS and ONE location in Bronx that has the `highest` number of collisions.((latitude, longitude) = (0,0) should not be chosen).

Save your result in a table named **`top_collisions`**.
The format of your table should be:


-----| borough  |   latitude | longitude |
-----|----------|------------|-----------|
 0   |  QUEENS  |  40.6807   |-73.8446429|  



```python
###
### YOUR CODE HERE
###

query = '''
    SELECT BOROUGH,LATITUDE	, LONGITUDE, Num_of_persons_injured
    from  nypd 
   
'''
df6=pd.read_sql_query(query, conn)
df7=df6[df6['BOROUGH']=='QUEENS']
df8=df6[df6['BOROUGH']=='BRONX']
df7=df7.groupby(['LATITUDE','LONGITUDE']).count()
df7.sort_values(by='Num_of_persons_injured',ascending=False, inplace=True)
df8=df8.groupby(['LATITUDE','LONGITUDE']).count()
df8.sort_values(by='Num_of_persons_injured',ascending=False, inplace=True)
dict1={'borough': 'QUEENS','latitude':df7.iloc[0].name[0],'longitude':df7.iloc[0].name[1]}
dict1
dict2={'borough': 'BRONX','latitude':df8.iloc[0].name[0],'longitude':df8.iloc[0].name[1]}
top_collisions=pd.DataFrame([dict1,dict2])


# Show your solution:
print(top_collisions)
top_collisions['latitude'].sum()
top_collisions['longitude'].sum()
```

Output:
```
borough   latitude  longitude
0  QUEENS  40.733497 -73.870369
1   BRONX  40.820189 -73.890675
-147.7610446
```

<!-- grade_id: exercise2, points: 3, grade: true, solution: false -->
```python
## Test Cell: Check the locations ##
assert (top_collisions['borough'].iloc[0]=='QUEENS')
assert (top_collisions['borough'].iloc[1]=='BRONX')
assert (top_collisions['latitude'].sum() == 81.55368630000001)
assert (top_collisions['longitude'].sum() == -147.76104459999999)
assert ((top_collisions['latitude']*top_collisions['longitude']).sum() == -6025.2298195065159)

print("\n(Passed!)")
```

Output:
```
(Passed!)
```

### Detail_injured table

Detail_injured table is required for the next question. It has 5 columns, displayed as below.

In terms of Database structures, the `Detail_injured` table's Unique_key column is a primary key, the same as the `NYPD` table's Unique_key column.
This means you can perform any SQL joins or Pandas merges between the two tables on this column.

```python
Detail_injured = pd.read_table('../resource/lib/publicdata/mt2nb9/Detail_injured.txt', sep=',')
print('Number of injured and killed:')
print('-------------------')
print(Detail_injured.head())
```

Output:
```
Number of injured and killed:
-------------------
   Unique_key  Num_of_persons_injured  Num_of_persons_killed  \
0     4090610                       0                      0   
1     4090611                       0                      0   
2     4090424                       0                      0   
3     4089381                       0                      0   
4     4089728                       0                      0   

   Num_of_pedestrians_injured  Num_of_pedestrians_killed  
0                           0                          0  
1                           0                          0  
2                           0                          0  
3                           0                          0  
4                           0                          0
```

```python
query = '''
    SELECT Unique_key,Zip_code,LATITUDE, LONGITUDE, Num_of_persons_injured
    from  nypd 
   
'''
df8=pd.read_sql_query(query, conn)
Detail_injured['injured']=Detail_injured['Num_of_persons_injured']+Detail_injured['Num_of_pedestrians_injured']
Detail_injured.sort_values(by='injured',ascending=False, inplace=True)
target=Detail_injured[['Unique_key','injured']]
df8.Unique_key=df8.Unique_key.astype(int)

Detail_injured.sort_values(by='injured',ascending=False, inplace=True)
target=Detail_injured[['Unique_key','injured']]
df8.Unique_key=df8.Unique_key.astype(int)


```

```python
df9=df8.merge(target,on='Unique_key',how='left')
df10=df9[['injured']].groupby([df9['LATITUDE'],df9['LONGITUDE']]).sum()
df10.sort_values(by='injured',ascending=False, inplace=True)
test=df10.iloc[1:11]
target2=[]
for i in range(10):
    result=dict()
    A=test.iloc[i].name[0]
    B=test.iloc[i].name[1]
    result['LATITUDE']=A
    result['LONGITUDE']=B
    result['Zip_code']=df9[(df9['LATITUDE']==A) & (df9['LONGITUDE']==B)]['Zip_code'].values[0]
    result['num_of_collisions']=len(df9[(df9['LATITUDE']==A) & (df9['LONGITUDE']==B)].index)
    target2.append(result)
target2
vic=pd.DataFrame(target2)
vic.sort_values(by='num_of_collisions',ascending=False, inplace=True)
vic

```

Output:
```
LATITUDE  LONGITUDE Zip_code  num_of_collisions
3  40.675736 -73.896853    11207                433
0  40.658578 -73.890623    11207                417
4  40.820189 -73.890675    10459                333
6  40.680015 -73.877428    11208                302
5  40.678163 -73.897477    11207                286
7  40.663183 -73.962449    11225                282
2  40.668798 -73.931120    11213                281
8  40.656160 -73.767353    11434                226
1  40.642058 -74.020420    11220                208
9  40.804376 -73.937420    10035                207
```

```python

```

**Exercise 3(3 points):** Use table **`nypd`** and table **`detail_injured`**. Both tables have primary key "Unique_key". 

Find the top 10 locations (by Latitude and Longitude) that have the largest number of injured people, which is the sum of `person_injured` and `pedestrians_injured` in table **`detail_injured`**. Please be aware it is not the same as `Num_of_persons_injured` in table **`nypd`**.

Report the locations (by Latitude and Longitude), zip code and the count of collisions for these top 10 locations, by the descending order of num_of_collisions.  

Save your result in a table named **`top_ten_locations`**. ((latitude, longitude) = (0,0) should not be chosen)

The format of your table should be:

-----| LATITUDE | LONGITUDE  | Zip_code |num_of_collisions|
-----|----------|------------|----------|-----------------|
 0   |  40.6807 | -73.8446429|  11200   |    10          |

```python
###
### YOUR CODE HERE
###
query = '''
    SELECT Unique_key,Zip_code,LATITUDE, LONGITUDE, Num_of_persons_injured
    from  nypd 
   
'''
df8=pd.read_sql_query(query, conn)
Detail_injured['injured']=Detail_injured['Num_of_persons_injured']+Detail_injured['Num_of_pedestrians_injured']
Detail_injured.sort_values(by='injured',ascending=False, inplace=True)
target=Detail_injured[['Unique_key','injured']]
df8.Unique_key=df8.Unique_key.astype(int)

Detail_injured.sort_values(by='injured',ascending=False, inplace=True)
target=Detail_injured[['Unique_key','injured']]
df8.Unique_key=df8.Unique_key.astype(int)

df9=df8.merge(target,on='Unique_key',how='left')
df10=df9[['injured']].groupby([df9['LATITUDE'],df9['LONGITUDE']]).sum()
df10.sort_values(by='injured',ascending=False, inplace=True)
test=df10.iloc[1:11]
target2=[]
for i in range(10):
    result=dict()
    A=test.iloc[i].name[0]
    B=test.iloc[i].name[1]
    result['LATITUDE']=A
    result['LONGITUDE']=B
    result['Zip_code']=df9[(df9['LATITUDE']==A) & (df9['LONGITUDE']==B)]['Zip_code'].values[0]
    result['num_of_collisions']=len(df9[(df9['LATITUDE']==A) & (df9['LONGITUDE']==B)].index)
    target2.append(result)
target2
vic=pd.DataFrame(target2)
vic.sort_values(by='num_of_collisions',ascending=False, inplace=True)
top_ten_locations=vic



print(top_ten_locations)
```

Output:
```
LATITUDE  LONGITUDE Zip_code  num_of_collisions
3  40.675736 -73.896853    11207                433
0  40.658578 -73.890623    11207                417
4  40.820189 -73.890675    10459                333
6  40.680015 -73.877428    11208                302
5  40.678163 -73.897477    11207                286
7  40.663183 -73.962449    11225                282
2  40.668798 -73.931120    11213                281
8  40.656160 -73.767353    11434                226
1  40.642058 -74.020420    11220                208
9  40.804376 -73.937420    10035                207
```

<!-- grade_id: exercise3, points: 3, grade: true, solution: false -->
```python
## Test Cell: Check the locations ##
assert top_ten_locations['num_of_collisions'].sum() == 2975 or top_ten_locations['num_of_collisions'].sum() == 2974
assert top_ten_locations[top_ten_locations['Zip_code'] == '11207']['num_of_collisions'].sum() == 1136
assert top_ten_locations.iloc[6]['LATITUDE'] == 40.66879779999999
assert top_ten_locations.iloc[6]['LONGITUDE'] == -73.93112009999999  
assert top_ten_locations.iloc[3]['LONGITUDE'] == -73.8774282
assert top_ten_locations.iloc[3]['LATITUDE'] == 40.680015000000004
assert top_ten_locations.iloc[7]['Zip_code'] == '11434'

print("\n(Passed!)")
```

Output:
```
(Passed!)
```

```python
# Some cleanup code
conn.close()
```

** Note: ** You've reached the end of this problem. Don't forget to restart the kernel and run the entire notebook from top-to-bottom to make sure you did everything correctly. If that is working, try submitting this problem. (Recall that you *must* submit and pass the autograder to get credit for your work.)